---
title: Test
description: Test
---

## An exercise title written in sentence case

```yaml
type: NormalExercise
key: e867f47227
lang: python
xp: 100
skills: 2
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
